% exec '/Users/blanchet/misc/E-weights2/PROVER/eproof' --tstp-in --tstp-out -l5 --split-clauses=4 --split-reuse-defs --simul-paramod --forward-context-sr --destructive-er-aggressive --destructive-er --presat-simplify --prefer-initial-clauses -tKBO6 -winvfreqrank -c1 -Ginvfreqconjmax -F1 --delete-bad-limit=150000000 -WSelectMaxLComplexAvoidPosPred -H'(4*FunWeight(SimulateSOS, 20,20,1.5,1.5,1,bit0:0,bit1:0,even_2_even:0,number_number_of_nat:0,pls:0,suc:8,zero_zero_nat:25),3*ConjectureGeneralSymbolWeight(PreferNonGoals,200,100,200,50,50,1,100,1.5,1.5,1),1*Clauseweight(PreferProcessed,1,1,1),1*FIFOWeight(PreferProcessed))' -tAutoDev --silent --cpu-limit=30 '/Users/blanchet/.isabelle/prob_e_1' 2>&1
% 2011-06-03 11:21:28
# Problem is unsatisfiable (or provable), constructing proof object
# SZS status Theorem
# SZS output start CNFRefutation.
fof(1, axiom,![X1]:(even_2_even(X1)=>even_2_even(suc(suc(X1)))),file('/Users/blanchet/.isabelle/prob_e_1', fact_0_even_Ointros_I2_J)).
fof(2, axiom,suc(suc(number_number_of_nat(bit0(bit1(pls)))))=number_number_of_nat(bit0(bit0(bit1(pls)))),file('/Users/blanchet/.isabelle/prob_e_1', fact_1_Suc__Suc__2__eq__4)).
fof(3, axiom,suc(suc(zero_zero_nat))=number_number_of_nat(bit0(bit1(pls))),file('/Users/blanchet/.isabelle/prob_e_1', fact_2_Suc__Suc__0__eq__2)).
fof(4, axiom,even_2_even(zero_zero_nat),file('/Users/blanchet/.isabelle/prob_e_1', fact_3_even_Ointros_I1_J)).
fof(5, conjecture,(even_2_even(number_number_of_nat(bit0(bit1(pls))))&even_2_even(number_number_of_nat(bit0(bit0(bit1(pls)))))),file('/Users/blanchet/.isabelle/prob_e_1', conj_0)).
fof(6, negated_conjecture,~((even_2_even(number_number_of_nat(bit0(bit1(pls))))&even_2_even(number_number_of_nat(bit0(bit0(bit1(pls))))))),inference(assume_negation,[status(cth)],[5])).
fof(7, plain,![X1]:(~(even_2_even(X1))|even_2_even(suc(suc(X1)))),inference(fof_nnf,[status(thm)],[1])).
fof(8, plain,![X2]:(~(even_2_even(X2))|even_2_even(suc(suc(X2)))),inference(variable_rename,[status(thm)],[7])).
cnf(9,plain,(even_2_even(suc(suc(X1)))|~even_2_even(X1)),inference(split_conjunct,[status(thm)],[8])).
cnf(10,plain,(suc(suc(number_number_of_nat(bit0(bit1(pls)))))=number_number_of_nat(bit0(bit0(bit1(pls))))),inference(split_conjunct,[status(thm)],[2])).
cnf(11,plain,(suc(suc(zero_zero_nat))=number_number_of_nat(bit0(bit1(pls)))),inference(split_conjunct,[status(thm)],[3])).
cnf(12,plain,(even_2_even(zero_zero_nat)),inference(split_conjunct,[status(thm)],[4])).
fof(13, negated_conjecture,(~(even_2_even(number_number_of_nat(bit0(bit1(pls)))))|~(even_2_even(number_number_of_nat(bit0(bit0(bit1(pls))))))),inference(fof_nnf,[status(thm)],[6])).
cnf(14,negated_conjecture,(~even_2_even(number_number_of_nat(bit0(bit0(bit1(pls)))))|~even_2_even(number_number_of_nat(bit0(bit1(pls))))),inference(split_conjunct,[status(thm)],[13])).
cnf(15,negated_conjecture,(~even_2_even(number_number_of_nat(bit0(bit0(bit1(pls)))))|~even_2_even(suc(suc(zero_zero_nat)))),inference(rw,[status(thm)],[14,11,theory(equality)])).
cnf(16,plain,(suc(suc(suc(suc(zero_zero_nat))))=number_number_of_nat(bit0(bit0(bit1(pls))))),inference(rw,[status(thm)],[10,11,theory(equality)])).
cnf(17,negated_conjecture,(~even_2_even(suc(suc(suc(suc(zero_zero_nat)))))|~even_2_even(suc(suc(zero_zero_nat)))),inference(rw,[status(thm)],[15,16,theory(equality)])).
cnf(18,negated_conjecture,(~even_2_even(suc(suc(zero_zero_nat)))),inference(csr,[status(thm)],[17,9])).
cnf(19,negated_conjecture,(~even_2_even(zero_zero_nat)),inference(spm,[status(thm)],[18,9,theory(equality)])).
cnf(20,negated_conjecture,($false),inference(rw,[status(thm)],[19,12,theory(equality)])).
cnf(21,negated_conjecture,($false),inference(cn,[status(thm)],[20,theory(equality)])).
cnf(22,negated_conjecture,($false),21,['proof']).
# SZS output end CNFRefutation
